Anti-inflammatory therapy with nebulised dornase alfa for severe COVID-19 pneumonia

Author response:

The following is the authors’ response to the previous reviews.

Suggestions to the authors:

• Please re-analyze findings by omitting from all Tables and Figures all data of comparators who were not randomized (BAC). I understand the difficulties of running this trial but the results of excess reduction of mortality do not allow the publication of a trial where comparators do not come from the randomized patient population.

We wish to thank the editors and reviewers for their useful comments. Given that the study was designed with both randomised and CC participants we can’t easily exclude the CC analysis from the paper. However, we do provide graphs for both randomised only and randomised and CC participants for the primary and secondary endpoints. The fact that the primary endpoint (CRP) results are mirrored in both instances is also informative form a trial design perspective and indicative of the effect of dornase alfa therapy on inflammation being robust enough to yield the same results with small and larger cohorts.

We agree that there are potential drawbacks of using contemporary controls. To address these potential biases we used CC patients recruited at the same time period at single site using the same selection criteria as the randomised group, which minimised potential bias. However, the enrolment and comparison of CRP in CC-BAC participants to concurrent randomised control R-BAC patients indicated that the two groups responded to BAC treatment in the same manner (Table 2, LS means log(CRP) 3.78 vs 3.53, P=0.386), whereas the R-BAC+DA vs R-BAC group comparison yielded significant differences (Table 2, LS means log(CRP) 3.1 vs 3.59, P=0.041). These comparisons mitigate to a large degree these potential problems.

Still, to make easy to distinguish the groups we now use the following unique nomenclature throughout the manuscript which is clearly defined on ln. 111 and state that comparisons of treated participants were performed with both control groups separately and combined.

R-BAC: Randomised BAC CC-BAC: Contemporary control BAC R-BAC+DA : Randomised BAC+ dornase alfa T-BAC: R-BAC + CC-BAC

In fact, the most important bias in our study, might actually be the placebo effect, given that participants randomised to BAC did not receive a nebulized control substance. We now discuss these points in more detail in the manuscript and modified the title by removing the reference to a randomised trial and clinical outcomes.

• The presentation remains confusing and the manuscript should be critically revised for clarity. There is a repetition of methods (e.g. lines 176-187 repeat 160-175) and redundant results (e.g. Figure S2, Table 3).

We apologise for the repetition. We removed the repeated text in the Exclusion criteria (lines 176-187 in the old manuscript).

Figure S2 is not related to Table 3. Figure S2 depicts baseline characteristics, whereas Table 3 complements the graph in Figure 3A but lists the mean daily value of the primary endpoint as requested by Reviewer 1 in the first round of revision.

At Table 4: the authors should select one method of illustration for lab results, either Table or figure, without repetitions

We agree and have removed Table 4 leaving the graphs instead.

• Regarding inclusion criteria, it is unclear whether high radiological suspicion is sufficient for inclusion or whether PCR based confirmation is required in all instances (differences in wording between lines 153 and 191), and under which oxygen requirements (lines 155 and 192)

We thank the reviewer for pointing this out. Indeed, radiological suspicion was not sufficient and all participants in this study had a positive PCR test as part of their diagnosis prior to inclusion in the study. The entire eligibility section was rewritten to reflect this important point.

• Table 1 should be merged with Table S2 and a better description of cohort baseline severity (P/F, SOFA, APACHE, organ support, number of patients in each point of the WHO severity score) and treatments should be made available.

We thank the reviewer for this suggestion. We have now merged Table 1 and S2 and included WHO ordinal severity information in Table 1, with median, average, SD, min and max values which reflect the participant distribution. Unfortunately, although the additional requested information was recorded, it was not systematically collected for the analysis of the trial and it was not straight forward to compile at this stage.

Changes from the earlier version to the revised version

Abstract

Abstract

Background

Cell-free (cf)-DNA, from cellular sources, including Neutrophil Extracellular Traps (NETs), is found in the circulation of COVID-19 patients and may contribute to immune dysregulation. This study determined whether pulmonary administration of the endonuclease, dornase alfa, reduced systemic inflammation by degrading local and systemic cf-DNA.

Methods

Earlier version: Eligible patients were randomized (3:1) to receive twice-daily nebulised dornase alfa in addition to best available care (BAC) or BAC alone for seven days or until discharge. A 2:1 ratio of matched contemporary controls (CC) provided additional comparators. The primary endpoint was improvement in C-reactive protein (CRP) over time, analysed using a repeated-measures mixed model, adjusted for baseline factors.

Revised version: Eligible patients were randomized (3:1) to receive twice-daily nebulised dornase alfa in addition to best available care (BAC) or BAC alone for seven days or until discharge. A 2:1 ratio of matched contemporary controls (CC-BAC) provided additional comparators. The primary endpoint was improvement in C-reactive protein (CRP) over time, analysed using a repeated-measures mixed model, adjusted for baseline factors.

Results

Earlier version: Between June 2020-October 2021 we recruited 39 evaluable patients: 30 randomised to dornase alfa; 9 randomised to BAC; with 60 CC. Dornase alfa reduced CRP by 33% compared to BAC. Least squares (LS) mean post-dexamethasone CRP fell from 101.9mg/L to 23.23 mg/L in the BAC+ dornase alfa group versus a fall from 99.5mg/L to 34.82 mg/L in the BAC group at 7 days; P=0.01. This effect of dornase alfa on CRP was confirmed with subgroup and sensitivity analyses. Dornase alfa increased the chance of live discharge by 63% (HR 1.63, 95% CI 1.01 to 2.61, P=0.03), increased lymphocyte counts (LSmean 1.08 vs 0.87, P=0.02) and reduced circulating cf-DNA and the coagulopathy marker D-dimer (LS mean: 570.78 vs 1656.96 μg/mL, P=0.004). Dornase alfa was well-tolerated.

Revised version: Between June 2020-October 2021 we recruited 39 evaluable patients: 30 randomised to dornase alfa (R-BAC+DA); 9 randomised to BAC (R-BAC); with the addition of 60 CC-BAC participants. Dornase alfa was well tolerated and reduced CRP by 33% compared to combined BAC groups (T-BAC). Least squares (LS) mean post-dexamethasone CRP fell from 101.9mg/L to 23.23 mg/L in the BAC+ dornase alfa group versus a fall from 99.5mg/L to 34.82 mg/L in the T-BAC group at 7 days; P=0.01. This effect of dornase alfa on CRP was confirmed with subgroup and sensitivity analyses that mitigated potential biases associated with the use of the CC-BAC group. Dornase alfa increased the chance of live discharge by 63% (HR 1.63, 95% CI 1.01 to 2.61, P=0.03), increased lymphocyte counts (LS mean: 1.08 vs 0.87, P=0.02) and reduced circulating cf-DNA and the coagulopathy marker D-dimer (LS mean: 570.78 vs 1656.96 μg/mL, P=0.004).

Added in the revised version: Conclusion

Removed in the revised version: Conclusions

Earlier version: We provide proof-of-concept evidence that dornase alfa reduces pathogenic inflammation in hospitalised patients with COVID-19 pneumonia.

Revised version: We provide proof-of-concept evidence that dornase alfa reduces pathogenic inflammation in hospitalised patients with COVID-19 pneumonia, suggesting that best available care can be improved by the inclusion of anti-inflammatory treatments that target damage-associated molecules.

Introduction

Earlier version: SARS-CoV-2 pneumonia can lead to hyperinflammation, coagulopathy, respiratory failure, and death (Siddiqi and Mehra, 2020; Zhou et al., 2020), even with less pathogenic variants and in immunised patients (2000). In severe COVID-19 neutrophil activation drives neutrophil extracellular trap (NET) formation (Radermecker et al., 2020; Zuo et al., 2020). NETs, are composed of DNA-histone complexes and have been associated with coagulopathy and endothelial dysfunction (Papayannopoulos, 2018). Moreover, extracellular histones promote inflammation, immune dysfunction and lethality in sepsis (Ioannou et al., 2022; Tsourouktsoglou et al., 2020; Xu et al., 2011; Xu et al., 2009). Digestion of chromatin DNA by endonucleases suppresses the proinflammatory activity of histones and enables their clearance from the circulation (Ioannou et al., 2022; Tsourouktsoglou et al., 2020). Consistently, DNAse I treatment reduces pathology in murine pulmonary viral infections (Cortjens et al., 2018; Pillai et al., 2016). Endogenous DNAse activity and NET clearance capacity are defective in severe COVID-19 pneumonia (Aramburu et al., 2022b), and the extent of these defects correlates with mortality. Hence, supplementation with exogenous DNases may facilitate extracellular chromatin degradation and increase survival.

Revised version: SARS-CoV-2 pneumonia can lead to hyperinflammation, coagulopathy, respiratory failure, and death (Siddiqi and Mehra, 2020; Zhou et al., 2020), even with less pathogenic variants and in immunised patients (2000). In severe COVID-19 neutrophil activation drives neutrophil extracellular trap (NET) formation (Radermecker et al., 2020; Zuo et al., 2020). NETs, are composed of DNA-histone complexes and have been associated with coagulopathy and endothelial dysfunction (Papayannopoulos, 2018). Moreover, extracellular histones promote inflammation, immune dysfunction and lethality in sepsis (Ioannou et al., 2022; Tsourouktsoglou et al., 2020; Xu et al., 2011; Xu et al., 2009). Digestion of chromatin DNA by endonucleases suppresses the proinflammatory activity of histones and enables their clearance from the circulation (Ioannou et al., 2022; Tsourouktsoglou et al., 2020). Consistently, DNAse I treatment reduces pathology in murine pulmonary viral infections (Cortjens et al., 2018; Pillai et al., 2016). Endogenous DNAse activity and NET clearance capacity are defective in severe COVID-19 pneumonia, and the extent of these defects correlates with mortality (Aramburu et al., 2022). Hence, supplementation with exogenous DNases may facilitate extracellular chromatin degradation and increase survival.

Pulmozyme, dornase alfa, is a recombinant human DNAase approved since 1993 for patients with cystic fibrosis (CF) (Konstan and Ratjen, 2012; Lazarus and Wagener†, 2019). Dornase alfa solubilizes NETs, reduces inflammation and improves pulmonary function in chronic and acute exacerbations of CF (Konstan and Ratjen, 2012; Papayannopoulos et al., 2011). Pulmozyme is safe and well-tolerated in children and adults with CF at doses up to 10mg BD. Nebulized pulmozyme does not increase circulating endonucleases. Hence, it remained unclear whether the degradation of extracellular chromatin in the lungs would reduce circulating chromatin and systemic hyper-inflammation.

Earlier version: To probe the effectiveness of DNase I as an anti-inflammatory therapy in severe COVID-19 pneumonia, we measured as our primary endpoint C-reactive protein (CRP). CRP is a reliable marker of systemic inflammation in severe infection(McArdle et al., 2004). In patients infected with COVID-19, a high CRP of ≥40 mg/L carried a poor prognosis with a lower cut-off of ≥35 mg/L in the elderly (Villoteau et al., 2021), particularly when combined with lymphopenia and coagulopathy marked by elevated D-dimer (Fisher et al., 2022; Liu et al., 2020; Luo et al., 2020; Smilowitz et al., 2021; Tornling et al., 2021; Ullah et al., 2020). Poor outcomes are associated with an early rise (Mueller et al., 2020), higher peak and delayed reduction in CRP (Cui et al., 2021). Moreover, elevated IL-6 and CRP can predict the need for mechanical ventilation(Herold et al., 2020). In addition, CRP correlated well with levels of DNAse and NET degradation activity in COVID-19 patient plasma suggesting a link between CRP and the capacity to degrade cf-DNA in the circulation (Aramburu et al., 2022b). CRP became a reliable primary endpoint in subsequent trials demonstrating the effectiveness of other systemic therapies such as namilumab or infliximab that target inflammatory cytokines (Fisher et al., 2022).

Revised version: To probe the effectiveness of DNase I as an anti-inflammatory therapy in severe COVID-19 pneumonia, we measured as our primary endpoint C-reactive protein (CRP). CRP is a reliable marker of systemic inflammation in severe infection (McArdle et al., 2004). In patients infected with COVID-19, a high CRP of ≥401mg/L carried a poor prognosis with a lower cut-off of ≥351mg/L in the elderly (Villoteau et al., 2021), particularly when combined with lymphopenia and coagulopathy marked by elevated D-dimer (Fisher et al., 2022; Liu et al., 2020; Luo et al., 2020; Smilowitz et al., 2021; Tornling et al., 2021; Ullah et al., 2020). Poor outcomes are associated with an early rise (Mueller et al., 2020), higher peak and delayed reduction in CRP (Cui et al., 2021). Moreover, elevated IL-6 and CRP can predict the need for mechanical ventilation(Herold et al., 2020). In addition, CRP correlated well with levels of DNAse and NET degradation activity in COVID-19 patient plasma suggesting a link between CRP and the capacity to degrade cf-DNA in the circulation (Aramburu et al., 2022). CRP became a reliable primary endpoint in subsequent trials demonstrating the effectiveness of other systemic therapies such as namilumab or infliximab that target inflammatory cytokines (Fisher et al., 2022).

Earlier version: In this proof-of-concept study we evaluated the effect of dornase alfa on inflammation, as measured by the impact on circulating CRP in patients with COVID-19, compared with best available care (BAC). The trial was initiated in June 2020 and was completed in October 2021. At the start of the trial only dexamethasone had been proven to benefit hospitalized COVID-19 pneumonia patients and was thus included in both arms of the trial. To increase the chance of reaching significance under challenging constraints in patient access, we opted to increase our sample size by using a combination of randomized individuals and available CRP data from matched contemporary controls (CC) hospitalized at UCL but not recruited to a trial. These approaches demonstrated that when combined with dexamethasone, nebulized DNase treatment was an effective anti-inflammatory treatment in randomized individuals with or without the implementation of CC data.

Revised version: In this proof-of-concept study we evaluated the effect of dornase alfa on inflammation, as measured by the impact on circulating CRP in patients with COVID-19, compared with best available care (BAC). The trial was initiated in June 2020 and was completed in October 2021. At the start of the trial only dexamethasone had been proven to benefit hospitalized COVID-19 pneumonia patients and was thus included in both arms of the trial. To increase the chance of reaching significance under challenging constraints in patient access, we opted to increase our sample size by using a combination of randomized individuals in the BAC group (R-BAC) and the dornase alfa treated arm (R-BAC+DA) with available CRP data from matched contemporary controls (CC-BAC) hospitalized at the same site at UCL but not recruited to a trial and using the same selection criteria as the randomised group to minimise potential biases. In the analysis framework we compared treated R-BAC-DA participants to both R-BAC and CC-BAC groups individually or combined (T-BAC). These approaches demonstrated that when combined with dexamethasone, nebulized DNase treatment was an effective anti-inflammatory treatment in randomized individuals with or without the implementation of CC-BAC data.

Methods

Sponsor and location

Earlier version: The trial was sponsored by University College London (UCL) and carried out at University College London Hospital UCLH with ethical (REC: 20/SC/0197, Protocol: 132333, RAS ID:283091) and UK MHRA approvals. All randomized participants provided informed consent. Consent for CC was covered by Health Service (Control of Participant Information) Regulations 2002. Safety and data integrity were overseen by the Trial Monitoring Group and Data Monitoring Committee. All data was collected at UCLH. Additional CRP data is reported from the Pa-COVID-19 study, Charité Universitätsmedizin Berlin with ethical approvaL, Berlin (EA2/066/20). Both Pa-COVID-19 and COVASE studies were carried out according to the Declaration of Helsinki and the principles of Good Clinical Practice (ICH 1996).

Revised version: The trial was sponsored by University College London (UCL) and carried out at University College London Hospital UCLH with ethical (REC: 20/SC/0197, Protocol: 132333, RAS ID:283091) and UK MHRA approvals. All randomized participants provided informed consent. Consent for CC-BAC participants was covered by Health Service (Control of Participant Information) Regulations 2002. Safety and data integrity were overseen by the Trial Monitoring Group and Data Monitoring Committee. All data was collected at UCLH. Additional CRP data is reported from the Pa-COVID-19 study, Charité Universitätsmedizin Berlin with ethical approval, Berlin (EA2/066/20). Both Pa-COVID-19 and COVASE studies were carried out according to the Declaration of Helsinki and the principles of Good Clinical Practice (ICH 1996).

Trial design

Earlier version: The COVASE trial was a single-site, randomised, controlled, parallel, open-label investigation. The primary endpoint was change in CRP. Screening was performed within 24 hours prior to the administration of dornase alfa (Figure 1A). Eligible, consented participants were randomly assigned (3:1) to either BAC plus nebulised dornase alfa (BAC+DA) or BAC alone. On Day 1 a baseline sample was collected. From Day 1-7 participants randomised to active arm received 2.5mg BD nebulised dornase alfa in addition to BAC. In all cases BAC included dexamethasone (6 mg/day) for 10 days or until discharge, whichever was shorter as per RECOVERY (Group et al., 2021). Participants received additional treatments at the discretion of their physicians. The primary analysis was performed on samples up to Day 7. The final trial visit occurred at day 35.

Revised version: The COVASE trial was a single-site, randomised, controlled, parallel, open-label investigation. The primary endpoint was change in CRP. Screening was performed within 24 hours prior to the administration of dornase alfa (Figure 1A). Eligible, consented participants were randomly assigned (3:1) to either BAC plus nebulised dornase alfa (R-BAC+DA) or R-BAC group alone. On Day 1 a baseline sample was collected. From Day 1-7 participants randomised to active arm received 2.5mg BD nebulised dornase alfa in addition to BAC. In all cases BAC included dexamethasone (6 mg/day) for 10 days or until discharge, whichever was shorter as per RECOVERY (Group et al., 2021). Participants received additional treatments at the discretion of their physicians. The primary analysis was performed on samples up to Day 7. The final trial visit occurred at day 35.

Trial design and Consort diagram

A. COVASE Trial Design.

Earlier version: B. Consort diagram summary. Numbers not in parentheses indicate the participants in the intention-to-treat population and the numbers in parentheses indicate the number of participants in the per-protocol population. A complete consort diagram is shown in Supplementary figure 1.

Revised version: B. Consort diagram summary. Numbers not in parentheses indicate the participants in the ITT population and the numbers in parentheses indicate the number of participants in the per-protocol population. A complete consort diagram is shown in Supplementary figure 1.

Earlier version: In addition to the recruited randomised group, we included contemporary controls (CC) to increase the sample size due to practical difficulties in recruiting patients in the UK. For every COVASE participant randomised to active treatment, 2 matched CC were included. CC were admitted to UCLH over the same time as randomised patients and treated with the same BAC, including dexamethasone. CCs fulfilled the inclusion and exclusion criteria of the COVASE study and were matched for age, gender, BMI, comorbidities and CRP. CRP matching was based either on admission CRP (CC1), or matched based on first CRP after starting dexamethasone (CC2).

Revised version: In addition to the recruited randomised group, we included CC-BAC participants to increase the sample size due to practical difficulties in recruiting patients in the UK. For every COVASE participant randomised to active treatment, 2 matched CC-BAC participants were included. CC-BAC participants were admitted to UCLH over the same time as randomised patients and treated with the same BAC, including dexamethasone. CC-BAC participants fulfilled the inclusion and exclusion criteria of the COVASE study and were matched for age, gender, BMI, comorbidities and CRP. CRP matching was based either on admission CRP, or matched based on the first CRP reading after starting dexamethasone.

Eligibility

Trial participant inclusion criteria

Added in the revised version: Male and female participants, aged ≥ 18 years, who were hospitalised for suspected Coronavirus (SARS-CoV)-2 infection confirmed by polymerase chain reaction (PCR) test and optionally with radiological confirmation with chest CT. Participants with stable oxygen saturation (>=94%) on supplementary oxygen and a CRP >= 30 mg/L. Participants that provided written informed consent to participate in the study and were able to comply with instructions and the use of a nebuliser.

Removed in the revised version: Male and female participants, aged ≥ 18 years Dornase alfa in COVID-19: Sponsor code: 132333 Protocol Version 2.0, 25.4.2020 Page 26 of 51

Removed in the revised version: Participants who are hospitalised for suspected Coronavirus (SARS-CoV)-2 infection confirmed by polymerase chain reaction (PCR) test or radiological confirmation with chest CT

Removed in the revised version: Participants with stable oxygen saturation (>=94%) on supplementary oxygen

Removed in the revised version: CRP >= 30 mg/L

Removed in the revised version: Participants will have given their written informed consent to participate in the study and are able to comply with instructions and nebuliser

Exclusion criteria

Added in the revised version: Females who were pregnant, planning pregnancy or breastfeeding. Concurrent and/or recent involvement in other research or use of another experimental investigational medicinal product that was likely to interfere with the study medication within 3 months of study enrolment. Serious condition meeting either respiratory distress with respiratory rate >=40 breaths/min, or an oxygen saturation<=93% on high-flow oxygen. Participants who required mechanical invasive or non-invasive ventilation at screening or had concurrent severe respiratory disease such as asthma, COPD and/or ILD, or any major disorder that in the opinion of the Investigator would interfere with the evaluation of the results or constitute a health risk for the trial participant. Participants with terminal disease and life expectancy <12 months without COVID-19, known allergies to dornase alfa and excipients, or participants who were unable to inhale or exhale orally throughout the entire nebulisation period.

Removed in the revised version: Females who are pregnant, planning pregnancy or breastfeeding

Removed in the revised version: Concurrent and/or recent involvement in other research or use of another experimental investigational medicinal product that is likely to interfere with the study medication within (specify time period e.g. last 3 months) of study enrolment

Removed in the revised version: Serious condition meeting one of the following:

Removed in the revised version: Respiratory distress with respiratory rate >=40 breaths/min

Removed in the revised version: oxygen saturation<=93% on high-flow oxygen

Removed in the revised version: Require mechanical invasive or non-invasive ventilation at screening

Removed in the revised version: Concurrent severe respiratory disease such as asthma, COPD and/or ILD

Removed in the revised version: Any major disorder that in the opinion of the Investigator would interfere with the evaluation of the results or constitute a health risk for the trial participant

Removed in the revised version: Terminal disease and life expectancy <12 months without COVID-19

Removed in the revised version: Known allergies to dornase alfa and excipients

Removed in the revised version: Participants who are unable to inhale or exhale orally throughout the entire nebulisation period

Removed in the revised version: So briefly

Removed in the revised version: Patients were excluded if they were:

Removed in the revised version: 1. pregnant, planning pregnancy or breastfeeding

Removed in the revised version: 2. Serious condition meeting one of the following:

Removed in the revised version: Respiratory distress with respiratory rate >=40 breaths/min

Removed in the revised version: oxygen saturation<=93% on high-flow oxygen

Removed in the revised version: 4. Require ventilation at screening

Removed in the revised version: 5. Concurrent severe respiratory disease such as asthma, COPD and/or ILD

Removed in the revised version: 6. Terminal disease and life expectancy <12 months without COVID-19

Removed in the revised version: 8. Known allergies to dornase alfa and excipients

Removed in the revised version: 9. Participants who are unable to inhale or exhale orally throughout the entire nebulisation period”

Participants

Earlier version: Adults (≥18 years of age) admitted to UCLH with confirmed SARS-Cov2 infection by RT-PCR and radiologically confirmed COVID-19 pneumonia on chest radiograph or CT-scan; oxygen saturation < 94% requiring supplemental oxygen; and evidence of hyperinflammation (CRP ≥ 30 mg/L, after administration of dexamethasone) were eligible. While all participants had CRP > 30 mg/L at screening, on two occasions the baseline CRP concentrations fell below

Revised version: Adults (≥18 years of age) admitted to UCLH with confirmed SARS-Cov2 infection by RT-PCR and radiologically confirmed COVID-19 pneumonia on chest radiograph or CT-scan; oxygen saturation < 94% requiring supplemental oxygen; and evidence of hyperinflammation (CRP ≥ 30 mg/L, after administration of dexamethasone) were eligible. While all participants had CRP > 30 mg/L at screening, on two occasions the baseline CRP concentrations fell below 30 mg/L after the patients had already agreed to participate. For full inclusion and exclusion criteria see Protocol in the Supplementary Material.

Removed in the revised version: 30 mg/L after the patients had already agreed to participate. For full inclusion and exclusion criteria see Protocol: Supplementary Appendix 1.

Outcomes

The primary outcome was the least square (LS) mean CRP up to 7 days or at hospital discharge whichever was sooner. Data beyond 7 days was included in the model to estimate the slope. Subsequently, the estimated slope in each group was used in order to calculate the LS mean at 7 days. Pre-specified secondary outcomes included days on oxygen; time to hospital discharge; mortality by day 35; and changes in clinically relevant biomarkers including lymphocyte count and D-dimer levels. Day 35 was chosen as being likely to include most early mortality due to COVID-19 being 4 weeks after completion of a week of treatment. (i.e. d7 of treatment +28 (4 x 7 days)).

Earlier version: Efficacy assessments of primary and secondary outcomes in the intention-to-treat (ITT) population were performed on all randomised participants who had received at least one dose of dornase alfa if randomized to treatment. For full details see Statistical Analysis Plan. The ITT was adjusted to mitigate the following protocol violations where one participant in the BAC arm and one in the DA arm withdrew before they received treatment and provided only a baseline CRP measurement available. The participant in the DA arm was replaced with an additional recruited patient. Exploratory endpoints were only available in randomised participants and not in the CC. In this case, a post hoc within group analysis was conducted to compare baseline and post-baseline measurements.

Revised version: Efficacy assessments of primary and secondary outcomes in the intention-to-treat (ITT) population were performed on all randomised participants who had received at least one dose of dornase alfa if randomized to treatment. For full details see Statistical Analysis Plan. The ITT was adjusted to mitigate the following protocol violations where one participant in the R-BAC arm and one in the R-BAC+DA arm withdrew before they received treatment and provided only a baseline CRP measurement available. The participant in the DA arm was replaced with an additional recruited patient. Exploratory endpoints were only available in randomised participants and not in the CC-BAC. In this case, a post hoc within group analysis was conducted to compare baseline and post-baseline measurements.

Sample size calculation

Size calculations were produced using the proc power function in SAS Version 9.4. These were conducted to achieve 80% power to detect difference in the active arm versus the control group at 5% level of significance. Based on a mean of 99mg/L in the control group and a common standard deviation of 62mg/L derived from the literature (Han et al., 2020; Zhou, 2020), a total sample size of 90 participants would provide sufficient power to detect a greater than a 40% relative difference for CRP in the dornase alfa group compared to the control group.

Earlier version: This study used existing data collected at UCLH from contemporary controls (CC) admitted with COVID-19. This gave a ratio of active versus comparator of 1:2. The required power, would result in 30 participants in the active treatment group and at least 60 in the control. An additional 10 participants were recruited as a control for exploratory objectives and to compare the characteristics of enrolled participants with CC. This gave a total of 40 participants enrolled in the study and 60 CC.

Revised version: This study used existing data collected at UCLH from CC-BAC participants admitted with COVID-19. This gave a ratio of active versus comparator of 1:2. The required power, would result in 30 participants in the active treatment group and at least 60 in the control group. An additional 10 radnomised control participants were recruited for exploratory objectives and to compare the characteristics of enrolled participants with CC-BAC participants. Therefore, a total of 40 participants were enrolled in the study and 60 CC-BAC participants were included.

Randomisation

A closed envelope method was implemented to randomize the participants into the control and active treatment groups.

Results

Patient charateristics

Earlier version: From June 2020-October 2021, 41 participants were recruited and randomised. 1 participant in BAC group was discharged before a second CRP measurement, and was excluded from all, except safety, analyses (CONSORT Diagram: Figure 1B and S1). One participant withdrew consent prior to receiving dornase alfa and was replaced and excluded from all analyses. 39 participants were included in the intention-to-treat analysis set, 30 BAC+DA and 9 BAC. One participant withdrew due to side-effects. This participant was removed from the per protocol population (PPP). All 39 participants were followed up for 35 days or death whichever was sooner. Two participants were excluded from the PPP. One from BAC as this was the only patient in whom randomisation occurred prior to dexamethasone being widely used in COVID-19 and a second participant who withdrew after one dose of dornase alfa (Figure 1B and S1). The trial ended when 40 eligible participants had been recruited, although thereafter 1 had participant in the BAC arm had to be excluded from the primary and secondary endpoint analysis.

Revised version: From June 2020-October 2021, 41 participants were recruited and randomised, with 1 participant in R-BAC group discharged before a second CRP measurement who was excluded from all analyses except for safety (CONSORT Diagram: Figure 1B and S1). One participant withdrew consent prior to receiving dornase alfa and was replaced and excluded from all analyses. 39 participants were included in the ITT analysis set, 30 R-BAC+DA and 9 R-BAC. One participant withdrew due to side-effects. This participant was removed from the per protocol population (PPP). All 39 participants were followed up for 35 days or death whichever was sooner. Two participants were excluded from the PPP and one from R-BAC as this was the only patient in whom randomisation occurred prior to dexamethasone being widely used in COVID-19 and a second participant who withdrew after one dose of dornase alfa (Figure 1B and S1). The trial ended when 40 eligible participants had been recruited, although thereafter 1 had participant in the R-BAC arm had to be excluded from the primary and secondary endpoint analysis.

Earlier version: Baseline characteristics were well balanced across groups (Table 1). Selection of 60 matched CC participants via propensity score matching was successful in ensuring the means of the characteristics included in propensity score matching were similar (Table 1 and Figure S2). The overall mean age was 56.8 years (56.8 years BAC+DA, 56.8 years BAC). The percentage of males was 75.8% (76.7% BAC+DA, 75.4% BAC). The most prevalent ethnicity was “White British”, with 30.3% of participants identifying in that category (33.3% BAC+DA, 29.0% BAC). The mean BMI was 28.0kg/m2 (27.8kg/m2 BAC+DA, 28.2kg/m2 BAC). The mean baseline CRP (post dexamethasone) was 100.2mg/L (101.9mg/L BAC=DA, 99.5mg/L) BAC. The proportion of participants with a comorbidity, defined as one or more of hypertension, diabetes, or cardiovascular disease, was 52.5% (46.7% BAC+DA, 55.1% BAC). All randomised participants except one received dexamethasone prior to randomisation, and 48 of the total 99 participants also received remdesivir or tocilizumab in addition to dexamethasone within the first 7 days. The average duration of dexamethasone treatment prior to dornase alfa was 1.13+/-0.79 days (Table S1). The last pre-dexamethasone CRP was also similar between groups, n mean of 125.0mg/L (me128.1mg/L BAC+DA, 122.7mg/L BAC). The number of days between dexamethasone initiation and baseline was 1.2 days (0.7 days BAC+DA, 1.3 days BAC). There were imbalances noted at baseline between the groups in white blood cell count, neutrophil count, procalcitonin count and D-dimer (Table S2).

Revised version: Baseline characteristics were well balanced across groups (Table 1). Selection of 60 matched CC-BAC participants via propensity score matching was successful in ensuring the means of the characteristics included in propensity score matching were similar (Table 1 and Figure S2). The overall mean age was 56.8 years (56.8 years R-BAC+DA, 56.8 years T-BAC). The percentage of males was 75.8% (76.7% R-BAC+DA, 75.4% T-BAC). The most prevalent ethnicity was “White British”, with 30.3% of participants identifying in that category (33.3% BAC+DA, 29.0% T-BAC). The mean BMI was 28.0kg/m2 (27.8kg/m2 R-BAC+DA, 28.2kg/m2 T-BAC). The mean baseline CRP (post dexamethasone) was 100.2mg/L (101.9mg/L R-BAC+DA, 99.5mg/L T-BAC). The proportion of participants with a comorbidity, defined as one or more of hypertension, diabetes, or cardiovascular disease, was 52.5% (46.7% R-BAC+DA, 55.1% T-BAC). All randomised participants, except one, received dexamethasone prior to randomisation. In addition, 48 of the total 99 participants also received remdesivir or tocilizumab in addition to dexamethasone within the first 7 days. The average duration of dexamethasone treatment prior to dornase alfa was 1.13+/- 0.79 days (Table S1). The last pre-dexamethasone CRP was also similar between groups with a mean of 125.0mg/L (128.1mg/L R-BAC+DA, 122.7mg/L T-BAC). The number of days between dexamethasone initiation and baseline was 1.2 days (0.7 days R-BAC+DA, 1.3 days T-BAC). There were imbalances noted at baseline between the groups in white blood cell count, neutrophil count, procalcitonin count and D-dimer (Table 1).

Clinical outcomes

Longitudinal CRP correlates with disease outcomes

Earlier version: To ascertain the clinical value of CRP as the COVASE trial’s primary endpoint, we first examined the correlation of CRP with mortality in a cohort of 63 hospitalized patients with severe COVID-19 (Aramburu et al., 2022b; Kurth et al., 2020; Messner et al., 2020). All participants who reached a WHO severity ordinance scale 7 were recruited at the Charité Hospital in Berlin in spring 2020 and did not receive anti-inflammatory therapies such as dexamethasone or anti-IL-6 antibody treatments. Plasma CRP measurements in 465 samples were segregated by disease outcome. This analysis revealed that deceased participants had significantly higher CRP than participants in the survivor group (Figure 2A). Clustering the 63 participant according to their longitudinal average CRP indicated that the frequency of mortality increased as the average CRP per patient increased (Figure 2B). Consistently, Mantel Cox survival analysis of the participants segregated into three groups with average longitudinal CRP ranges of 0-100 mg/L, 100-200 mg/L and 200-450 mg/L, showed that survival decreased significantly as the average CRP increased (Figure 2C). This results highlighted the strong correlation between CRP and survival in hospitalized individuals with severe COVID-19 pneumonia, suggesting that the blood CRP concentration was a strong predictor of clinical outcomes such as the probability of survival.

Revised version: To ascertain the clinical value of CRP as the COVASE trial’s primary endpoint, we first examined the correlation of CRP with mortality in a cohort of 63 hospitalized patients with severe COVID-19 (Aramburu et al., 2022; Kurth et al., 2020; Messner et al., 2020). All participants who reached a WHO severity ordinance scale 7 were recruited at the Charité Hospital in Berlin in spring 2020 and did not receive anti-inflammatory therapies such as dexamethasone or anti-IL-6 antibody treatments. Plasma CRP measurements in 465 samples were segregated by disease outcome. This analysis revealed that deceased participants had significantly higher CRP than participants in the survivor group (Figure 2A). Clustering the 63 participants according to their longitudinal average CRP indicated that the frequency of mortality increased as the average CRP per patient increased (Figure 2B). Consistently, Mantel Cox survival analysis of the participants segregated into three groups with average longitudinal CRP ranges of 0-100 mg/L, 100-200 mg/L and 200-450 mg/L, showed that survival decreased significantly as the average CRP increased (Figure 2C). These results highlighted the strong correlation between CRP and survival in hospitalized individuals with severe COVID-19 pneumonia, suggesting that the blood CRP concentration was a strong predictor of clinical outcomes such as the probability of survival.

Earlier version: Longitudinal CRP predicts survival probability in severe COVID-19 pneumonia. Data from the Berlin COVID-19 study

Revised version: Longitudinal CRP predicts survival probability in severe COVID-19 pneumonia.

Earlier version: A. Individual CRP concentrations in 465 plasmas from 63 participants with maximum WHO severity grade 7 COVID-19 pneumonia segregated into survivor and deceased groups.

Revised version: A. Individual CRP concentrations in 465 plasmas from 63 participants with maximum WHO severity grade 7 COVID-19 pneumonia segregated into survivor and deceased groups, extracted from the Berlin COVID-19 study.

B. Participants ordered by their longitudinal average CRP concentrations shown in the left column. Mortality is depicted in yellow in the right column.

C. Kaplan Meier survival probabilities (left panel) and numbers at risk (right panel) for patients segregated into three categories of longitudinal average CRP ranges: 0-100, 100-200, and 200-450 mg / L. Statistical significance (P), Hazard ratios (HR) and 95% confidence intervals (95% CI) for group 1 against group 3 and group 2 against group 3 are shown below the survival plot.

Statistics by Mann-Whitney and Mantel-Cox log rank tests.

Primary outcome

Earlier version: Next, we analysed the results from our COVASE participants whose individual CRP traces over time are shown in Figure S3A. Blood collection for both BAC and BAC+DA groups occurred at similar times and comparable frequencies (Figure S3B). For the ITT group, the LS mean log (CRP) over 7 days follow-up was 3.15 (95% confidence interval [CI] 2.87 to 3.42) BAC+DA (n=30), and 3.55 (95% CI, 3.35 to 3.75) in BAC (n=69) (Table 2, 3 and Figure 3A), p=0.01. This indicates a reduction in mean CRP of approximately 33% in the BAC+DA (23.23 mg/mL) compared to BAC (34.82 mg/mL) at mean follow-up over 7 days. This effect of dornase alfa on CRP was confirmed in various other subgroup analyses (Table 2): the per-modified-protocol population only; the participants who were randomised into the COVASE trial, excluding the CC (Figure 3B); participants who were randomised to BAC+DA in the COVASE trial, and CC excluding those randomised to BAC only. In addition, to ensure that the CC did not have a significantly different CRP trajectory to those randomised to BAC, we compared participants who were randomised to BAC with CC by excluding those randomised to BAC+DA and found no significant differences (Table 2). Sensitivity analyses supported the observed effect on CRP. These included the log(CRP) as an area under the curve; the CC matched for their last pre-dexamethasone CRP measurement as opposed to their first CRP after starting dexamethasone; and the effect of remdesivir or tocilizumab (Table 2).

Revised version: Next, we analysed the results from our COVASE participants. The individual CRP traces over time are shown in Figure S3A. Blood collection for both R-BAC and R-BAC+DA groups occurred at similar times and comparable frequencies (Figure S3B). In the participants who were randomised into the COVASE trial, excluding the CC-BAC participants the LS mean Ln (CRP) over 7 days follow-up was 3.10 (95% confidence interval [CI] 2.84 to 3.35) R-BAC+DA (n=30), and 3.59 (95% CI, 3.13 to 4.06) in R-BAC (n=9) (Table 2 and Figure 3A). Moreover, in the ITT group which included randomised and CC-BAC groups, the LS mean Ln (CRP) over 7 days follow-up was 3.15 (95% confidence interval [CI] 2.87 to 3.42) R-BAC+DA (n=30), and 3.55 (95% CI, 3.35 to 3.75) in T-BAC (n=69) (Table 2 and Figure 3A), p=0.01. This indicates a reduction in mean CRP of approximately 33% in the R-BAC+DA (23.23 mg/mL) compared to T-BAC (34.82 mg/mL) at mean follow-up over 7 days. The suppressive effect of dornase alfa on CRP was confirmed in various other subgroup analyses: the per-modified-protocol population only; participants who were randomised to BAC+DA in the COVASE trial, and CC-BAC participants alone excluding those randomised to BAC (Table 2). In addition, to ensure that the CC-BAC participants did not have a significantly different CRP trajectory to those randomised to BAC, we compared participants who were randomised to BAC (R-BAC) with the CC-BAC group by excluding those randomised to R-BAC+DA and found no significant differences (Table 2). Sensitivity analyses supported the observed effect on CRP. These included the Ln(CRP) as an area under the curve; the CC-BAC matched for the last pre-dexamethasone CRP measurement as opposed to their first CRP after starting dexamethasone; and the effect of remdesivir or tocilizumab (Table 2). A daily mean Ln CRP measurement chart is also provided (Table 3).

Analysis of primary and clinical endpoints

Earlier version: A and B. Fitted mean (95% confidence interval) from mixed model, with natural log (CRP) over 7 days follow-up as the outcome, adjusting for natural log baseline CRP, age, sex, BMI, serious comorbidity (Diabetes, Cardiovascular disease or hypertension), time and a treatment × time interaction. P-value generated by comparing least-square means between arms. (A) Intention to treat (ITT) population (Blue: CC and participants randomised to BAC, N=69; Pink: participants randomised to BAC+DA, N=30) and (B) randomised participants only: (Blue: participants randomised to BAC, N=9; Pink: participants randomised to BAC+DA, N=30).

Revised version: A. Fitted mean (95% confidence interval) from mixed model of natural log (CRP) over 7 days follow-up as the outcome. (Left panel) randomised participants only: Blue: participants randomised to R-BAC, N=9; Pink: participants randomised to R-BAC+DA, N=30. (Right panel) ITT population: Blue: T-BAC (CC-BAC and R-BAC) N=69; Pink: R-BAC+DA, N=30. Results were adjusted for natural log baseline CRP, age, sex, BMI, serious comorbidity (Diabetes, Cardiovascular disease or hypertension), time and a treatment × time interaction. P-value generated by comparing least-square means between arms.

Earlier version: C. Distribution of participants based on the change in CRP measured as a ratio of the final CRP reading within the 7 day treatment period over the baseline CRP reading per patient. Statistical analysis by Fisher’s test.

Revised version: B. Distribution of participants based on the change in CRP measured as a ratio of the final CRP reading within the 7-day treatment period over the baseline CRP reading per patient. Statistical analysis by Fisher’s test.

Earlier version: D. Kaplan-Meier plot showing time to discharge from hospital from baseline. ITT population. Hazard ratio from Cox proportional hazards model adjusted for baseline CRP, age, sex, BMI, serious comorbidity (diabetes, cardiovascular disease of hypertension). P-value from log-rank test. (Blue: CC and participants randomised to BAC, N=69. Pink: participants randomised to BAC+DA, N=30).

Revised version: C. Kaplan-Meier plot showing time to discharge from hospital from baseline. Hazard ratio from Cox proportional hazards model adjusted for baseline CRP, age, sex, BMI, serious comorbidity (diabetes, cardiovascular disease of hypertension). P-value from log-rank test. (Blue: CC-BAC and participants randomised to R-BAC, N=69. Pink: participants randomised to R-BAC+DA, N=30).

Earlier version: E. Kaplan-Meier plot showing time to death over 35 days follow up. ITT population. Hazard ratio from Cox proportional hazards model adjusted for baseline CRP, age, sex, BMI, serious comorbidity (Diabetes, Cardiovascular disease of hypertension). P-value from log-rank test. (Blue: CC and participants randomised to BAC, N=69. Pink: participants randomised to BAC+DA, N=30). Abbreviations: BAC-best available care, CRP-C-reactive protein, DA-dornase alfa, ITT-intention-to-treat.

Revised version: D. Kaplan-Meier plot showing time to death over 35 days follow up. Hazard ratio from Cox proportional hazards model adjusted for baseline CRP, age, sex, BMI, serious comorbidity (Diabetes, Cardiovascular disease of hypertension). P-value from log-rank test.

Earlier version: To better understand the effects of dornase alfa on the CRP of the treated population we also performed a frequency distribution analysis of the change in CRP over time. Instead of fitting a slope on the CRP values of the individual samples of each participant, we plotted the fraction between the CRP measurement of the baseline and final sample collected for each participant during the 7 day treatment period. We then plotted the distribution of participants within the different ranges of change in CRP (Figure 3C). Compared to BAC, the BAC+DA treatment group exhibited a distribution towards larger ranges in CRP reduction with better suppression of the fraction of patients that did not respond to treatment as indicated by an increase in CRP (fraction: 1 to 6). These results indicated that the addition of dornase alfa to dexamethasone cause a significant and sustained reduction in CRP in a greater number of COVID-19 infected participants.

Revised version: To better understand the effects of dornase alfa on the CRP of the treated population we also performed a frequency distribution analysis of the change in CRP over time. Instead of fitting a slope on the CRP values of the individual samples of each participant, we plotted the fraction between the CRP measurement of the baseline and final sample collected for each participant during the 7 day treatment period. We then plotted the distribution of participants within the different ranges of change in CRP (Figure 3B). Compared to T-BAC, the R-BAC+DA treatment group exhibited a higher CRP reduction distribution with fewer participants that did not respond to treatment as indicated by an increase in CRP concentrations. Therefore, the addition of dornase alfa to dexamethasone resulted in a significant and sustained reduction in CRP in a greater number of COVID-19-infected participants compared to treatment with dexamethasone alone.

Secondary outcomes

Earlier version: Length of hospitalisation was analysed as a time-to-event outcome of alive discharge from hospital censored at 35 days. The hazard ratio observed in the Cox proportional hazards model was 1.63 (95% CI, 1.01 to 2.61), p=0.03 (Table 4, Figure 3D and S3C). Showing that throughout 35 days follow-up, there was a 63% higher chance of discharge alive at any given time-point in BAC+DA compared to BAC. Although the rate of discharge was similar in 50% of patients, 80% discharge occurred by 8 days in BAC+DA whereas, whereas the same proportion was reached at 30 days in BAC. This trend was also seen when only the randomised participants were considered, although not powered to reach significance and with a smaller HR of 1.18 (95% CI, 0.52-2.69), p=0.62, (Table S3 and Figure S3D).

Revised version: The length of hospitalisation was analysed as a time-to-event outcome of alive discharge from the hospital censored at 35 days. The hazard ratio observed in the Cox proportional hazards model was 1.63 (95% CI, 1.01 to 2.61), p=0.03 (Figure 3C and S3C). Showing that throughout 35 days follow-up, there was a 63% higher chance of discharge alive at any given time-point in R-BAC+DA compared to T-BAC. Although the rate of discharge was similar in 50% of patients, 80% discharge occurred by 8 days in R-BAC+DA whereas, whereas the same proportion was reached at 30 days in T-BAC. This trend was also seen when only the randomised R-BAC participants were considered, although not powered to reach significance and with a smaller HR of 1.18 (95% CI, 0.52-2.69), p=0.62, (Table S2 and Figure S3C).

Removed in the revised version: Over 7 days of follow up there was no significant difference between BAC+DA versus BAC alone in either the fraction of participants admitted to ICU (23.3% versus 21.74%), p=0.866, or the length of ICU stay, LS mean 21.25 (95% CI, 4.65 to 37.84) hours versus 19.85 (95% CI, 8.00 to 31.70) hours, p=0.883. The same was seen over 35-day follow-up, LS mean 55.21 95% CI, −23.59 to 134.00) hours versus 60.60 (95% CI, 4.34 to 116.86) hours, p=0.905. At any point during the 35 days follow-up, 23% of BAC+DA were admitted to ICU compared to 23.19% BAC, p= 0.983 (Table S4).

Earlier version: There was no significant difference in time requiring oxygen between the two groups, at either 7 days, LS mean 94.32 (95% CI, 72.8 to 115.79) hours, versus 88.96 (95% CI, 73.64 to 104.29) hours, p=0.662, or 35 days, LS mean 133.22 (95% CI, 52.01 to 214.43) hours versus 156.35 (95% CI, 98.36, 214.33) hours, p=0.618. At 35 days follow up, there were only 9 randomised participants to evaluate, but mean oxygen use tends to a reduction of 123 hours with BAC+DA, versus 241 hours for BAC, p=0.187 (Table S4)

Revised version: Over 7 days of follow up there was no significant difference between R-BAC+DA and T-BAC in either the fraction of participants admitted to ICU (23.3% versus 21.74%), p=0.866, or the length of ICU stay, LS mean 21.25 (95% CI, 4.65 to 37.84) hours versus 19.85 (95% CI, 8.00 to 31.70) hours, p=0.883. The same was seen over 35-day follow-up, LS mean 55.21 95% CI, -23.59 to 134.00) hours versus 60.60 (95% CI, 4.34 to 116.86) hours, p=0.905. At any point during the 35 days follow-up, 23% of R-BAC+DA were admitted to ICU compared to 23.19% T-BAC, p= 0.983 (Table S3). Similarly, there was no significant difference in the time requiring oxygen between the two groups (R-BAC-DA vs T BAC), at either 7 days, LS mean 94.32 (95% CI, 72.8 to 115.79) hours, versus 88.96 (95% CI, 73.64 to 104.29) hours, p=0.662, or 35 days, LS mean 133.22 (95% CI, 52.01 to 214.43) hours versus 156.35 (95% CI, 98.36, 214.33) hours, p=0.618. At 35 days follow up, there were only 9 randomised participants to evaluate, but mean oxygen use exhibited a reduction of 123 hours in R-BAC+DA participants, versus 241 hours for the T-BAC group, p=0.187 (Table S3).

Earlier version: The time to event data was censored at 28 days post last dose (up to d35) for the randomised participants and at the date of the last electronic record for the CC. Over 35 days follow up, 1 person amongst the 30 patients in BAC+DA died, compared to 8 of the 69 participants in BAC. The hazard ratio observed in the Cox proportional hazards model was 0.47 (95% CI, 0.06 to 3.86), indicating a trend towards a reduced chance of death at any given time-point in BAC+DA compared to BAC, but this did not reach significance p= 0.460 (Figure 3E and S3D). The hazard ratio observed in the Cox proportional hazards model (95% CI) was 0.47 (0.06, 3.86), which estimates that throughout 35 days follow-up, there was a 53% reduced chance of death at any given timepoint in the BAC + dornase-alfa group compared to the BAC group, though the confidence intervals are wide due to a small number of events. The p-value from a log-rank test was 0.460, which does not reach statistical significance at an alpha of 0.05.

Revised version: The time to event data was censored at 28 days post last dose (up to d35) for the randomised participants (R-BAC and R-BAC+DA) and at the date of the last electronic record for the CC-BAC group. Over 35 days follow up, 1 person amongst the 30 patients in the R-BAC+DA group died, compared to 8 of the 69 T-BAC participants. The hazard ratio observed in the Cox proportional hazards model was 0.47 (95% CI, 0.06 to 3.86), indicating a trend towards a reduced chance of death at any given time-point in R-BAC+DA compared to T-BAC, but this did not reach significance p= 0.460 (Figure 3D and S3C). The hazard ratio observed in the Cox proportional hazards model (95% CI) was 0.47 (0.06, 3.86), which estimates that throughout 35 days follow-up, there was a 53% reduced chance of death at any given timepoint in the R-BAC+DA group compared to the T-BAC group. However, the confidence intervals are wide due to the small number of events and consequently the p-value from a log-rank test was 0.460, which does not reach statistical significance.

Earlier version: There was no significant difference at either 7-or 35-days follow-up, in the number of participants that required mechanical ventilation in BAC+DA compared with BAC (16.67% vs 13.04%), p=0.628. Amongst participants that were ventilated, the mean length of mechanical ventilation at 7 days follow-up in BAC+DA was 76.8 hours, compared to 88.78 in BAC. At 35 days follow-up, the mean length of mechanical ventilation in BAC+DA was 76.8 hours compared to 411.17 hours in BAC (Table S4). There was no significant difference in superadded bacterial pneumonia at either 7-or 35-days follow-up: 7 days, 1 (3.33%) participant in BAC+DA compared to 3 (4.35%) participants in BAC, p= 0.934; 35 days, 2 (6.67%) participants in BAC+DA had bacterial pneumonia, compared to 3 (4.35%) participants in BAC, p=0.548 (Table S4).

Revised version: There was no significant difference at either 7- or 35-days follow-up, in the number of participants that required mechanical ventilation in R-BAC+DA compared with T-BAC (16.67% vs 13.04%), p=0.628. Amongst participants that were ventilated, the mean length of mechanical ventilation at 7 days follow-up in R-BAC+DA particpants was 76.8 hours, compared to 88.78 in the T-BAC group. At 35 days follow-up, the mean length of mechanical ventilation in BAC+DA was 76.8 hours compared to 411.17 hours in T-BAC participants (Table S3). There was no significant difference in superadded bacterial pneumonia at either 7 or 35 days follow-up: 7 days, 1 (3.33%) participant in BAC+DA compared to 3 (4.35%) participants in T-BAC, p= 0.934; 35 days, 2 (6.67%) R-BAC+DA participants had bacterial pneumonia, compared to 3 (4.35%) T-BAC participants, p=0.548 (Table S3).

Earlier version: Blood analysis with no adjustment for multiple testing showed a significant treatment effect in BAC+DA vs. BAC group for three parameters: lymphocyte counts, D-dimer, and procalcitonin (PCT). First, BAC+DA exhibited higher lymphocyte counts with LS mean of 0.87 (95% CI, 0.76-0.98) in the BAC group vs. 1.08 (95% CI, 0.92-1.27) in BAC+DA, p=0.02 (Table 4 and Table S3). Patients with lymphopenia at baseline (<1×109 lymphocytes/L) exhibited a greater increase in blood lymphocyte numbers in BAC+DA than in BAC during the entire length of treatment (Figure 4A). Furthermore, D-dimer levels were lower in BAC+DA compared to BAC, with LS mean D-dimer difference of 1657 (95% CI, 3131-877) (Table 4, Table S3, Figure 4B and Figure S4A). Our analysis indicated lower levels of PCT in participants that received BAC+DA, mean 0.18 ng/mL (95% CI, −0.2-0.56) compared to BAC alone, mean 1.31 ng/mL (95% CI, 0.56-2.05), p=0.005 (Table 4). Repeat analysis excluding the CC population replicated these results, and changes in all 3 parameters were significant (Table S3).

Revised version: Blood analysis with no adjustment for multiple testing showed a significant treatment effect in the BAC+DA vs. T-BAC group comparison for three parameters: lymphocyte counts, D-dimer, and procalcitonin (PCT). First, R-BAC+DA exhibited higher lymphocyte counts with LS mean of 0.87 (95% CI, 0.76-0.98) in the T-BAC group vs. 1.08 (95% CI, 0.92-1.27) in R-BAC+DA participants, p=0.02 ((Figure 4A, Table S2). Amongst individuals with lymphopenia at baseline (<1×109 lymphocytes/L), the R-BAC+DA group exhibited a greater increase in blood lymphocyte numbers compared to the T-BAC group throughout the entire length of treatment. Furthermore, D-dimer levels were lower in R-BAC+DA participants compared to T-BAC participants, with LS mean D-dimer difference of 1657 (95% CI, 3131-877) ((Figure 4B and S4A). R-BAC+DA participants also exhibited lower levels of PCT, mean 0.18 ng/mL (95% CI, -0.2-0.56) compared to the T-BAC group, mean 1.31 ng/mL (95% CI, 0.56-2.05), p=0.005. Repeat analysis excluding the CC-BAC population replicated these results, and changes in all 3 parameters were significant (Table S2).

Analysis of secondary and exploratory endpoints in blood

Earlier version: A. Difference between the lymphocyte count for each day of the treatment period and the baseline in each patient who exhibited lymphopenia at baseline (<1×109 lymphocytes/mL). Mean and 95%CI interval is shown with statistical analysis by two-way Anova.

Revised version: A. Difference between the lymphocyte count for each day of the treatment period and the baseline in each ITT participant who exhibited lymphopenia at baseline (<1×109 lymphocytes/mL). Mean and 95%CI interval is shown with statistical analysis by two-way Anova.

Earlier version: B. Mean D-dimer levels per day in randomised BAC (blue) and BAC+DA (pink) participants with error bars depicting 95% CI in randomised BAC (blue) and BAC+DA (DA) participants (pink). Statistical difference by by mixed effects Anova analysis.

Revised version: B. Mean D-dimer levels per day in randomised R-BAC (blue) and R-BAC+DA (pink) participants with error bars depicting 95% CI. Statistical difference by mixed effects Anova analysis.

Earlier version: C. Mean cf-DNA levels per day in randomised BAC (blue) and BAC+DA (pink) participants, with error bars depicting standard deviation. Statistical analysis by mixed effects Anova.

Revised version: C. Mean cf-DNA levels per day in randomised R-BAC (blue) and R-BAC+DA (pink) participants, with error bars depicting standard deviation. Statistical analysis by mixed effects Anova.

Earlier version: D. Correlation between the final cf-DNA levels and ratio of CRP at day-7 normalized to the baseline CRP (CRPfinal/CRPbaseline) per participant. Fitting by non-linear regression.

Revised version: D. Correlation between the final cf-DNA levels and ratio of CRP at day-7 normalized to the baseline CRP (CRPfinal/CRPbaseline) per randomised participant. Fitting by non-linear regression.

Earlier version: E. Correlation between D-dimer and cell-free (cf) DNA levels in the blood of participants randomised to BAC (blue) or to BAC+DA (DA) (pink), where samples have been segregated depending on whether the corresponding levels of cf-DNA were below or above 100 μg/mL. Statistical analysis by unpaired parametric t-test.

Revised version: E. Correlation between D-dimer and cf-DNA levels in the blood of participants randomised to R-BAC (blue) or to R-BAC+DA (DA) (pink), where samples were segregated depending on whether the corresponding levels of cf-DNA were below or above 100 μg/mL. Statistical analysis by unpaired parametric t-test.

Exploratory outcomes

Earlier version: Given the role of circulating cf-DNA in pathology, we examined whether the pulmonary administration of dornase alfa influenced systemic cf-DNA levels. Compared to the plasma of anonymized healthy donor volunteers at the Francis Crick institute (HD), plasma cf-DNA levels were elevated in both BAC and DA-treated COVASE participants. There was no difference in baseline plasma cf-DNA levels between the two groups. However, during the treatment period cf-DNA was reduced in participants randomised to BAC+DA compared to BAC alone (Figure 4C and Figure S4B). There was a positive correlation between the levels of cf-DNA in the final sample collected during the treatment period and the ratio of final/baseline CRP (CRPfinal/CRPbaseline) (Figure 4D). Moreover, samples with cf-DNA >100μg/mL had significantly higher D-dimer levels compared to samples with cf-DNA levels <100μg/mL (Figure 4E). Hence, dornase alfa treatment was associated with a suppression of circulating cf-DNA that correlated with a more sustained reduction in CRP.

Revised version: Given the potential role of circulating cf-DNA in pathology, we examined whether the pulmonary administration of dornase alfa influenced cf-DNA levels circulating in plasma. Compared to anonymized healthy donor volunteers at the Francis Crick institute (HD), cf-DNA levels were elevated in the plasmas of all R-BAC and R-BAC-DA randomised COVASE participants. There was no difference in baseline plasma cf-DNA levels between the two randomised participant groups. Notably, during the treatment period plasma cf-DNA was reduced in R-BAC+DA compared to R-BAC participants ((Figure 4C and Figure S4B). Moreover, there was a positive correlation between the levels of cf-DNA in the final sample collected during the treatment period and the magnitude of CRP reduction measured as the ratio of final over baseline CRP (CRPfinal/CRPbaseline) ((Figure 4D). Moreover, samples with cf-DNA >100μg/mL exhibited significantly higher D-dimer levels compared to samples with cf-DNA levels <100μg/mL ((Figure 4E). Hence, dornase alfa treatment was associated with a suppression of circulating cf-DNA that correlated with a reduced risk for coagulopathy and a more sustained reduction in CRP.

Safety

Earlier version: Dornase alfa was well tolerated with no systemic effects, consistent with its short half-life (Table S5). There were 10 reported adverse events (AEs) in 9 participants with BAC versus 30 AEs in 30 participants with BAC+DA (Table S6). Of these, one was reported by the clinical team as definitely related and one as unlikely to be related to study drug (Table S5). No treatment-related serious AEs were reported.

Revised version: Dornase alfa was well tolerated with no systemic side-effects which is consistent with its short half-life in vivo (Table S4). Overall, there were 10 reported adverse events (AEs) in the study, in 9 R-BAC and 30 R-BAC+DA participants (Table S5). Of these, one was reported by the clinical team as definitely related and one as unlikely to be related to dornase alfa. No treatment-related serious AEs were reported.

Discussion

Earlier version: Our findings show that nebulised dornase alfa significantly reduced CRP in patients with severe COVID-19 pneumonia and systemic inflammation even when already receiving dexamethasone. This was robust with several sensitivity analyses including analysis of only randomised participants (i.e excluding CC). The secondary outcomes were consistent with the same directionality. Despite not being formally powered for length of stay, dornase alfa reduced time to discharge over 35 days.

Revised version: Our findings show that nebulised dornase alfa significantly reduced systemic inflammation as measured by plasma CRP concentrations in patients with severe COVID-19 pneumonia even when already receiving dexamethasone. Dornase alfa exhibited a robust anti-inflammatory effect as assessed with several sensitivity analyses including analysis of only randomised participants, excluding CC-BAC participants. These findings confirm the usefulness of CRP as a sensitive readout to monitor the impact of anti-inflammatory agents in small cohorts. Improvements in care were also indicated by other biomarkers and secondary endpoints when both randomised and CC-BAC participants were considered. Moreover, despite not being formally powered for length of stay, dornase alfa reduced time to discharge over 35 days.

Earlier version: Identification of novel/repurposed treatments effective for COVID-19 were hampered by patient recruitment to competing studies during a pandemic. This resulted in small studies with inconclusive or contrary findings. Our study design took advantage of frequent repeated measures of CRP in each patient, to allow a smaller sample size to determine efficacy/futility. A similar CRP-based approach was applied in CATALYST (Fisher et al., 2022) and ATTRACT studies (Tornling et al., 2021). In addition, we used contemporary controls as additional comparators to efficiently use limited resources.

Revised version: During the pandemic, the identification of novel and repurposed treatments effective for COVID-19 was hampered by patient recruitment to competing studies during pandemic. This resulted in small studies with inconclusive or contrary findings. Our study design took advantage of frequent repeated measures of CRP in each patient, to allow a smaller sample size to determine efficacy. CRP was also a primary readout in CATALYST (Fisher et al., 2022) and ATTRACT studies (Tornling et al., 2021). In addition, we used contemporary controls as additional comparators to use limited resources more efficiently.

Earlier version: Following the RECOVERY trial, dexamethasone became standard care in patients with COVID-19 pneumonitis requiring oxygen. We recruited participants with CRP ≥ 30 mg/L, on the day after receiving dexamethasone to minimise steroid-dependent effects on CRP. The finding that dornase alfa can significantly reduce CRP (by average 33%) in participants receiving dexamethasone suggests a complementary mechanism of action. Dornase alfa may provide a treatment choice for patients with mild COVID-19 pneumonia not requiring oxygen, in whom dexamethasone may be harmful(Group et al., 2021).

Revised version: Following the RECOVERY trial, dexamethasone became the standard care in patients with COVID-19 pneumonitis requiring oxygen. We recruited participants with CRP ≥ 30 mg/L, on the day after receiving dexamethasone to minimise steroid-dependent effects on CRP. The finding that dornase alfa can significantly reduce CRP in participants receiving dexamethasone suggests a complementary mechanism of action with significant improvement in existing standard anti-inflammatory care. Moreover, dornase alfa may be a treatment choice for patients with mild COVID-19 pneumonia not requiring oxygen, in whom dexamethasone may be harmful (Group et al., 2021).

Earlier version: Our primary effect was consistently reflected in secondary outcomes. Dornase alfa increased the chance of live discharge by 63% at any time up to 35 days, with a trend to reduce hospital stay. This impact on hospital occupancy during a pandemic cannot be overestimated. In addition, dornase alfa significantly increased lymphocyte counts and reduced D-dimer. CRP levels are associated with venous thromboembolic disease in COVID-19 with the worst outcomes seen in patients with high CRP and D-dimers (Smilowitz et al., 2021). In a metanalysis of 32 studies involving 10 491 COVID-19 patients, elevated CRP (OR 4.37, p<0.00001), lymphopenia (3.33, p<0.00001), elevated D-dimer (3.39 (p<0.00001) and elevated PCT (6.33, p<0.00001) were independent markers of poor outcomes(Malik et al., 2021).

Revised version: The primary endpoint effect was consistently reflected in secondary outcomes. Dornase alfa increased the chance of live discharge by 63% at any time up to 35 days, reduce overall hospital stay. The reduction in hospital occupancy during the COVID-19 pandemic proved critical in sustaining life-saving treatment availability. In addition, dornase alfa significantly increased lymphocyte counts and reduced D-dimer. In addition to the increased mortality data we provide here, CRP levels are associated with venous thromboembolic disease in COVID-19 with the worst outcomes seen in patients with high CRP and D-dimers (Smilowitz et al., 2021). In a metanalysis of 32 studies involving 101491 COVID-19 patients, elevated CRP (OR 4.37, p<0.00001), lymphopenia (3.33, p<0.00001), elevated D-dimer (3.39 (p<0.00001) and elevated PCT (6.33, p<0.00001) were independent markers of poor outcomes (Malik et al., 2021).

Earlier version: Our study design offered a solution to the early screening of compounds for inclusion in larger platform trials. The study took advantage of frequent repeated measures of quantifiable CRP in each patient, to allow a smaller sample size to determine efficacy/futility than if powered on clinical outcomes. We applied a CRP-based approach that was similar to the CATALYST and ATTRACT studies (Fisher et al., 2022; Tornling et al., 2021). CATALYST showed in much smaller groups (usual care, 54, namilumab, 57 and infliximab, 35) that namilumab that is an antibody that blocks the cytokine GM-CSF reduced CRP even in participants treated with dexamethasone whereas infliximab that targets TNF-α had no significant effect on CRP. This led to a suggestion that namilumab should be considered as an agent to be prioritised for further investigation in the RECOVERY trial. A direct comparison of our results with CATALYST is difficult due to the different nature of the modelling employed in the two studies. In general, Dornase alfa exhibited comparable significance in the reduction in CRP compared to standard of care as described for namilumab at a fraction of the cost. Furthermore, endonuclease therapies may prove superior to cytokine blocking monotherapies, as they are unlikely to increase the risk for microbial co-infections that have been reported for antibody therapies that neutralize cytokines that are critical for immune defence such as IL-1β, IL-6 or GM-CSF.

Revised version: Our study design offered a solution to the early screening of compounds for inclusion in larger platform trials. The study took advantage of frequent repeated measures of quantifiable CRP in each patient, to enable the determination of efficacy using a smaller sample size to than if powered were based on clinical outcomes. We applied a CRP-based approach that was similar to the CATALYST and ATTRACT studies (Fisher et al., 2022; Tornling et al., 2021). CATALYST showed in much smaller groups (usual care, 54, namilumab, 57 and infliximab, 35) that the GM-CSF-blocking antibody namilumab reduced CRP even in participants treated with dexamethasone whereas infliximab that targets TNF-α had no significant effect on CRP. This led to a suggestion that namilumab should be considered as an agent to be prioritised for further investigation in the RECOVERY trial. A direct comparison of our results with CATALYST is difficult due to the different nature of the modelling employed in the two studies. Nevertheless, dornase alfa exhibited comparable significance in the reduction in CRP compared to standard of care as described for namilumab at a fraction of the cost. Furthermore, endonuclease therapies may prove safer and more applicable than cytokine-blocking monotherapies, as they are unlikely to increase the risk for microbial co-infections which are a risk factor associated with antibody therapies that neutralize inflammatory cytokines that are critical for immune defence such as IL-1β, IL-6 or GM-CSF.

Earlier version: In addition, Dornase alfa reduced circulating cf-DNA levels, suggesting that this nebulized enzyme exerts systemic effects on circulating chromatin and inflammatory mediators. The reduction in cf-DNA suggests that by stripping the DNA from chromatin, dornase alfa suppresses the proinflammatory properties of histones and potentiates their degradation(Aramburu et al., 2022a; Papayannopoulos et al., 2011). The inverse correlation between final circulating cf-DNA and change in CRP during treatment highlights the link between circulating chromatin and systemic inflammation. The correlation of d-dimer with cf-DNA and the reduction in d-dimer with dornase alfa is consistent with the reported pro-thrombotic role of NETs in the alveoli (Radermecker et al., 2020). Moreover, the difference in the recovery rates for lymphopenia is in line with a link between extracellular chromatin and lymphocyte death in sepsis (Ioannou et al., 2021). Another reason why recombinant endonuclease treatment may be warranted in patients with acute infections is the reduction in endogenous chromatin degradation capacity that correlates with a high risk for mortality (Aramburu et al., 2022a).

Revised version: In addition, dornase alfa reduced circulating cf-DNA levels, suggesting that lung administration of the enzyme exerts systemic effects on circulating chromatin. The reduction in plasma cf-DNA levels suggests that by stripping the DNA from chromatin, dornase alfa suppresses the proinflammatory properties of histones and potentiates their degradation (Papayannopoulos et al., 2011; Tsourouktsoglou et al., 2020). The inverse correlation between final circulating cf-DNA and the reduction in CRP during treatment is consistent with link between circulating chromatin and systemic inflammation observed in animal models (Ioannou et al., 2022). The correlation between d-dimer and cf-DNA as well as the reduction in D-dimer after dornase alfa treatment is consistent with the reported pro-thrombotic role of NETs in the alveoli (Radermecker et al., 2020). Moreover, the improvement in the recovery rates of lymphopenia support the link between extracellular chromatin and lymphocyte death during sepsis that we also detected in individuals with severe COVID-19 pneumonia (Aramburu et al., 2022; Ioannou et al., 2022). Another reason why recombinant endonuclease treatment may be warranted in severe COVID-19 and microbial sepsis infections is the significant reduction in endogenous NET chromatin degradation capacity that strongly correlates with a high risk for mortality in these conditions (Aramburu et al., 2022).

Earlier version: Nebulized dornase alfa treatment has several advantages. Whilst immunisation has reduced COVID-19 hospital admissions, a substantial population remain unvaccinated or immunocompromised. Hence, there will always be a need for virally agnostic therapies that retain efficacy as viruses mutate. Nebulised dornase alfa controls immune pathology in a range of infections. Moreover, it can be administered safely and effectively outside the health-care setting. Three other small trials of dornase alfa (totalling 18 patients) in COVID-19 have reported improved oxygenation (Holliday et al., 2021; Okur et al., 2020; Weber et al., 2020). One small study indicated improvements in plasma and sputum proteomic profiles (Fisher et al., 2021). Despite differing study designs, patient populations and endpoints, the overall trend suggests an improvement in clinical outcomes.

Revised version: Nebulized dornase alfa treatment has several advantages. Whilst immunisation has reduced COVID-19 hospital admissions, viral evolution and immune escape may still cause substantial mortality and other symports associated with long-COVID. Hence, there will always be a need for virally agnostic therapies that retain efficacy as viruses mutate. Importantly, nebulised dornase alfa has the potential to control immune pathology in a range of infections. Moreover, it can be administered safely and effectively in its nebulized form, outside the health-care setting. Three other small trials of dornase alfa (totalling 18 patients) in COVID-19 have reported improved oxygenation (Holliday et al., 2021; Okur et al., 2020; Weber et al., 2020). One small study indicated improvements in plasma and sputum proteomic profiles (Fisher et al., 2021). Despite differing study designs, patient populations and endpoints, there is a consensus for improvement in clinical outcomes.

Earlier version: There are limitations in this study. This single centre open-label study was designed to powerfully report futility or efficacy despite including just under 40 patients and 60 CC. However, the trial was not powered to report mortality nor to overcome confounders, such as the use of antivirals and tocilizumab, an IL-6 inhibitor recognised to reduce CRP (Galvan-Roman et al., 2021), or the impact of an open-label study on influencing other therapeutic/ discharge decisions. Although underpowered, we demonstrate a trend to a reduction in CRP with dornase alfa in participants that had received tocilizumab and/or remdesivir, and those that had not. Moreover, the open-label nature of the study could potentially introduce bias. Nevertheless, we tried to minimize this by applying a standard testing schedule and discharge criteria.

Revised version: Yet, there are a few limitations in this study. This single centre open-label study was designed to powerfully report futility or efficacy despite including just under 40 randomised patients and 60 CC-BAC. However, the trial was not powered to report mortality nor to overcome confounders, such as the use of antivirals and tocilizumab, an IL-6 inhibitor recognised to reduce CRP (Galvan-Roman et al., 2021), or the impact of an open-label study on influencing other therapeutic/ discharge decisions. Although underpowered, we demonstrate a trend to a reduction in CRP with dornase alfa in participants that had received tocilizumab and/or remdesivir, and those that had not. Moreover, the open-label nature of the study could potentially introduce a placebo effect bias. Nevertheless, we tried to minimize this by applying a standard testing schedule and discharge criteria.

Earlier version: In conclusion, we demonstrate that nebulised dornase alfa significantly reduces inflammation in hospitalised patients with severe COVID-19 pneumonia leading to improved clinical profiles and earlier discharge from hospital. These results confirm the role of cf-DNA in promoting systemic inflammation, coagulopathy and immune dysfunction in acute respiratory infections and suggest that recombinant DNases provide an additional mode of anti-inflammatory action that can broaden the effectiveness and sustainability of anti-inflammatory regimens. These encouraging data warrant further investigation in other respiratory infections characterized by hyperinflammation in the setting of cell lysis and/ or NET formation.

Revised version: In conclusion, we demonstrate that nebulised dornase alfa significantly reduces inflammation in hospitalised patients with severe COVID-19 pneumonia leading to improved clinical profiles and earlier discharge from hospital. These results confirm the role of cf-DNA in promoting systemic inflammation, coagulopathy and immune dysfunction in acute respiratory infections and suggest that recombinant endonucleases provide an additional mode of action that can broaden the effectiveness and sustainability of anti-inflammatory regimens. These encouraging data warrant further investigation in other serious respiratory infections characterized by hyperinflammation, cell lysis and NET formation.

Supporting information